Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "MUHAMMAD ABDURREHMAN ASIF"
COLLABORATORS = ""

```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):

    start_index = 0   #starting index
    end_index = 0   #ending index
    maximum = 0   #initializing the sum to be 0 so that we can replace it
    
    for i in range(len(A)+1):  #iterates over the entire length of the array
        for j in range(len(A)+1):  #we want this to iterate over the entire array as well
            current = sum(A[i:j])  #creates the 'maximum' subarray by adding the i'th and j'th element
            if current > maximum:   #compares the current subarray to any previously large ones
                maximum = current   #replaces it if it is bigger
                start_index = i     #replaces the start_index with i value for the max subarray
                end_index = j-1     #to avoid overcounting, we must subtract 1 to get the ending index
    return (start_index,end_index,maximum)

bruteforce_max_subarray([-2,-3,4,-1,-2,1])
```

Output:
```
(2, 2, 4)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Computing this problem mentally would be pretty easy, however, due to the nature of the sort of algorithm we are currently running, we would need to rerun this algorithm.For the first input array,we found the maximum sum to be the individual element '4'. All other elements are negative excluding the last one,which is smaller than 4.Therefore,the computation would look at 4 being the best possible answer. Once we add 8 at the end, mentally it would be easier to see that (1+8) > 4,however, for the computer to figure this out, the algorithm would need to run again.

If the computer had outputted the last element to be the maximum subarray,or be the endpoint of the maximum subarray,then it would be easy to assume that any positive integer appended to the array would be a part of the maximum subarray since it would positively affect it.However, this is not the case here.

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False, however there is a certain nuance to this answer.

If the maximum subarray in A is different than the maximum subarray in B, we need to run the algorithm again to see how much of an impact 8 has. Only rerunning it would prove to us which one is greater. If max of A > max of B (assuming we include the 8),then the 8 will not have an impact (for example if the maximum subarray of A is just the element 10).On the flipside, since we need to be careful about the fact that the sum of max B can become greater by appending 8 as we have observed in Question 2. 

Generally speaking however,a maximum subarray should only be deduced once all the elements are considered and having different maximums does not mean appending an element would not be important.


## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    if len(x) <= 1: #ensures that it returns the first element if the array only has one element
        return(x[0])

    maximum = float('-inf')    #initializing the maximum as a sentinel so each sum is greater
    for i in range(len(x)):    #runs the for loop
        maximum= sum(x[i:len(x)])   
        
        if maximum > mx:   #replaces the maximum sum if a greater one is computed
            mx = maximum
    return(mx)
    
        
 

B = [-2-3,4,-1,-2,1,8]
incremental_max_subarray(B,4)
```

Output:
```
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    mx = float("-inf")  #adding the sentinel so that every maximum is greater than negative infinity
    
    for i in range(len(A)):  #running iteration over every subarray
        mx = incremental_max_subarray(A[0:i+1], mx)  #replaces the maximum subarray by one which is bigger
    
    return mx

L = [-2,-4,6,-5,3,6,2,7,-4,3]
K = [i for i in range (0,10)]

assert(max_subarray(L) == 19)

assert(max_subarray(K) == 45) #cumulative sum of all elements

```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
The divide and conquer algorithm should be more efficient. The incremental subarray still runs iterations and has to run 2 iterations for each subarray. First, it needs to iterate over each single element of the array,and then it needs to iterate for each subarray and compute when there is increments happening. Thus it should have a runtime efficiency of O(N^2). 
On the other hand, we have already established that the runtime efficiency of divide and conquer algorithms is O(NlogN). This is more efficient and so over a larger array it will take much lesser time to compute.

```python

```