Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "MUHAMMAD ABDURREHMAN ASIF"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    
    #checks if maxheap is empty or if the element is lesser than the rootnode
    if not maxh or elem <= maxh[0] * -1:
        heapq.heappush(maxh, elem * -1) #heappush into maxheap
        
    #if the value is greater than root
    else:
        heapq.heappush(minh, elem) #heappush to rootnode
    
    
    #check which heap is larger and send the root to the other half
    if len(minh) - len(maxh) >1:
        heapq.heappush(maxh, heapq.heappop(minh) * -1)
    
    elif len(maxh) - len(minh) >1:
        heapq.heappush(mih, heapq.heappop(maxh) * -1)


```

## Question 2 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    
    #equal heaps
    if len(minh) == len(maxh):
        return ((maxh[0]*-1+minh[0])/2)
    
    #if min heap is greater
    elif len(minh) > len(maxh):
        return minh[0]
    
    #if maxheap is greater
    else:
        return maxh[0] * -1
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
2.0
3
4.0
5
6.0
7
8.0
9
10.0
11
12.0
13
14.0
15
16.0
17
18.0
19
20.0
21
22.0
23
24.0
25
26.0
27
28.0
29
30.0
31
32.0
33
34.0
35
36.0
37
38.0
39
40.0
41
42.0
43
44.0
45
46.0
47
48.0
49
50.0
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The add_to_median_heap performs two kinds of tasks. One are comparisons, and the second is heappush. Comparisons will take up an O(1) time complexity,whereas heappush takes O(logn). Since the worst case scenario would require us to iterate over all elements and keep running heappush on them,the worst case complexity would be O(logn)

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
This will always be O(1) since we do not perform any iterations, any loops, or any other functions that are computationally expensive. We perform a simple comparison and on the basis of that, we only remove the root nodes.

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
Both methods have their own benefits and cons. The vanilla way provides us with something that is easy to visualize and intuitively understand. A sorted list, selecting or computing the exact center of that list is something everyone can conceptualize. Doing it using python depends on the sort of algo we use -- first we would need to find an efficient sorting algo, then we would need to find an inexpensive median finder and combine both. Thus conventional sorting algos such as insertion sort or quicksort need to be adjusted appropriately for efficiency.

On the other hand, the algorithm we created now requires a deeper and more meaningful understanding of how heaps work and how we can use the root nodes to calculate the median. The plus point of this method is that it is computationally efficient, providing us with the median in O(nlogn) efficiency -- ( O(nlogn) because we run logn, n times)

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst,k):
    quick_sort(lst,0,len(lst)-1) 
    return(lst[k])    
    
    
def partition(A,p,r):
    """
    Assume r<len(A) and p>=0
    """
    x = A[r]
    i = p - 1
    
    for j in range(p,r):
        if A[j] <= x:
            i += 1
            A[i], A[j] = A[j], A[i]
        
    A[i+1], A[r] = A[r], A[i+1]
    return i + 1

def quick_sort(A,p,r):
    if p<r:
        q = partition(A,p,r)
        quick_sort(A,p,q-1)
        quick_sort(A,q+1,r)
    return A

```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))
```

Output:
```
0
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
There are 2 relations that can be made here. Either, the kth element is the median, or the kth element is the first or last element (the same we discussed as a problem for the quicksort before). 

Since we are only interested in the half of the list before or after the partition that contains our element of interest, at each stage we will break the function down by 2. Thus, the best case would take a T(n/2) + O(N) time. O(N) would be the comparisons complexity and T(N/2) for the partition and breaking down of the list.

The worst case would mean breaking down the list until only the first or last element remain -- therefore a runtime of T(N-1). We must perform the comparisons as well, and thus it would be T(n-1) + O(N)

## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
T(N) = T(N/2) + O(N)
T(k) = (T(N/4) + O(N/2)) + N ..... 
we can continually break this down but we have encountered functions like this before, and we see a pattern being created  (O(N/2^n)) which would sum to O(2N) since it is a geometric progression converging. O(2N) can be rewritten as O(N)

## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
T(N) = T(N-1) + O(N)
T(N) = T(N-2) + O(N-1) + O(N) .....
T(k) = O(N(n-1)). In the worst case, we would get O(N^2)

```python

```